<a href="https://colab.research.google.com/github/Carolinsrainbow/UC_AnalisisdeDatosEnPython/blob/main/Laboratorio4.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Laboratorio 4 (Trabajo final, parte uno)

Bienvenido al tercer laboratorio del curso Análisis de Datos en Python.

Escribe el código correspondiente para contestar cada una de las preguntas. Fíjate muy bien en el nombre de las variables, ya que de lo contrario, tu trabajo no podrá ser evaluado de manera correcta.

Debes ir ejecutando las celdas de codigo en el orden que aparecen, para evitar tener errores.

El laboratorio consta de dos partes: la primera, que realizarás en este jupyter notebook y la segunda, que se encuentra en el recurso llamado "Trabajo final (parte 2)"

```python
## Primero importamos las librerias y los datos que utilizaremos.
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
import datetime
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import StandardScaler
from sklearn.preprocessing import MinMaxScaler
from sklearn.linear_model import LinearRegression
from sklearn.metrics import r2_score
from sklearn.metrics import mean_squared_error
from sklearn.metrics import mean_absolute_error
from sklearn.cluster import KMeans
from sklearn import preprocessing
import matplotlib.pyplot as plt
import warnings
warnings.filterwarnings('ignore')
data = pd.read_excel('Tarea 4.xlsx', sheet_name = 'cellphones data')
ratings = pd.read_excel('Tarea 4.xlsx', sheet_name = 'cellphones ratings')
users = pd.read_excel('Tarea 4.xlsx', sheet_name = 'cellphones users')
```

### Pregunta 1

Utilice el método de agrupación “K-Means” y construya 3 agrupaciones de celulares, considerando los siguientes atributos

`['internal memory','RAM','performance','main camera','selfie camera','battery size']`

¿Qué puede decir sobre los precios de venta de los celulares de cada grupo?

Para responder a lo anterior, construye un dataframe que contega los 3 clusters y para cada grupo obtenga el precio promedio, la desviación estandar del precio, precio mínimo y precio máximo.

**MUY IMPORTANTE**
* NO debes aproximar ni redondear ningún valor
* Utilizar la función `preprocessing.MinMaxScaler()` para estandarizar datos.
* Usa la función `KMeans` con `random_state=42`. Los datos que entregas a esta función deben estar estandarizados, según el punto anterior.

Debes crear una tabla de la forma:


|   | Grupo k means | Precio promedio | Desv est precio | Precio mínimo | Precio máximo |
|---|---|---|---|---|---|
|0 | 0  | XXXXX  |  XXXXX  |  XXXXX  |  XXXXX  |
|1 | 1  | XXXXX  |  XXXXX  |  XXXXX  |  XXXXX  |
|2 | 2  | XXXXX  |  XXXXX  |  XXXXX  |  XXXXX  |

Utiliza los mismos nombres para los Headers: Grupo k means, Precio promedio, Desv est precio, Precio mínimo, Precio máximo. Es importante que respetes las mayúsculas de dichos nombres tal cual son señalados.
Debe guardarse en la variable 'data_grupos'

<!-- grade_id: cell-90bc898769b7238d, points: 0, grade: false, solution: true -->
```python
data_grupos=None

columnas = ['internal memory','RAM','performance','main camera','selfie camera','battery size']

# your code here
raise NotImplementedError
```

Haremos un test visible para ver si estamos por buen camino, usaremos de ejemplo el primer grupo (Grupo 0). Tu tabla debe comenzar de la siguiente forma:

|   | Grupo k means | Precio promedio | Desv est precio | Precio mínimo | Precio máximo |
|---|---|---|---|---|---|
|0 | 0  | 225.375000 |  122.057290  |  129  |  500  |


<!-- grade_id: cell-4065dead0f031387, points: 10, grade: true, solution: false -->
```python
# Revisión del primer cluster

data_cluster_0 = data_grupos[data_grupos["Grupo k means"] == 0]

assert np.round(data_cluster_0["Precio promedio"].values,1) == 225.4, "Promedio del cluster 0 con errores"
assert np.round(data_cluster_0["Desv est precio"].values, 1) == 122.1, "Desviación estandar del cluster 0 con errores"

assert np.round(data_cluster_0["Precio mínimo"].values, 1) == 129.0, "Mínimo del cluster 0 con errores"
assert np.round(data_cluster_0["Precio máximo"].values, 1) == 500, "Máximo del cluster 0 con errores"

```

Ahora haremos el test real, no debes modificar ni escribir nada en la siguiente celda, solo debes ejecutarla.

<!-- grade_id: cell-09925935c47b0412, points: 10, grade: true, solution: false -->
```python


```

### Pregunta 2
Calcule la correlación (bivariada) entre el precio y todas las variables cuantitativas con información sobre los celulares. ¿Cuáles son las 3 variables que presentan una mayor correlación con el precio?
Para responder a la pregunta, debes completar las variables variable_1, variable_2 y variable_3 con las 3 variables que poseen mayor correlacion, siendo variable_1 la que mayor correlación posee, variable_2, la segunda mayor y variable_3 la tercera mayor.

Ejemplo: si la variable 'main camera' fuera la variable con mayor correlación. Entonces deben dejar
variable_1 = 'main camera', si la segunda mayor fuera 'RAM' entonces variable_2 = 'RAM' (deben respetar el nombre con sus espacios y mayúsculas)

NOTA: Deben generar las correlaciones de la forma que más cómodos se sientan, no tendrán puntaje aquellos que completan las variables sin entregar un respaldo con las correlaciones calculadas.

<!-- grade_id: cell-0b2dc7e6cd810e1b, points: 0, grade: false, solution: true -->
```python
variable_1=None
variable_2=None
variable_3=None

# your code here
raise NotImplementedError

```

Revisemos que los nombres de las variables hayan sido bien escritos,

```python
for v in [variable_1, variable_2, variable_3]:
    print(v in data.columns)
```

<!-- grade_id: cell-4188e1d2f8026918, points: 20, grade: true, solution: false -->
```python

```

### Pregunta 3
Construya una regresión multivariada entre precio y las 3 variables determinadas en la parte (2). Utilice un 20% de los datos para testear y un random_state de 42. El entorno de este jupyter notebook permite el LinearRegression(normalize=True), así que úsenlo.
Una vez construido su modelo, debe almacenar la siguiente información:

    'coeficientes': para esto utilicen el comando .coef_ (nombredesumodelo.coef_) y lo almacenan directamente en la variable sin redondear ni aproximar nada.
    'intercepto': para esto utilicen el comando .intercept_ (nombredesumodelo.intercept_) y lo almacenan directamente en la variable sin redondear ni aproximar nada.
    'r2_score_final': también sin redondear ni aproximar nada.
    
Adicionalmente, genera una gráfica que contraste los valores reales con la predicción.


<!-- grade_id: cell-520335fb547c3200, points: 0, grade: false, solution: true -->
```python
coeficientes=None
intercepto=None
r2_score_final=None

# your code here
raise NotImplementedError

```

<!-- grade_id: cell-e7ea26b98e7417a9, points: 20, grade: true, solution: false -->
```python


```

### Pregunta 4
Construya una agrupación, utilizando K-Means, considerando ranking promedio (puesto por el comprador), edad y sexo. Escoja, con el método del codo, la cantidad de grupos. ¿Qué puede decir de las similitudes o diferencias intra e inter grupos?
Utilice el min_max_scaler para la estandarización de los datos y utilice un random_state de 42 en la creación de Kmeans.
Hint: para el método del códogo, haga un range hasta 15.

Para responder esta pregunta, complete la variable 'clusters' con el número de clusters que obtenga mediante el método del codo, si su resultado es 10, entonces 'clusters'=10. (como int).
Elija el segundo codo, el primero será muy notorio y evitende, pero el segundo será la opción correcta.

<!-- grade_id: cell-2efe7df1270d370c, points: 0, grade: false, solution: true -->
```python
#En esta primera parte, calcule los clusters mediante el método del codo.
clusters=None

# your code here
raise NotImplementedError



```

```python
# Revisemos que la variable haya quedado correctamente asignada a tu respuesta,
print(clusters)
```

Ahora se procederá a evaluar los clusters obtenidos, no modifique ni altere la celda que está acontinuación, solo debe ejecutarla.

<!-- grade_id: cell-fba53982f0704ebe, points: 10, grade: true, solution: false -->
```python

```

Ahora debe aplicar el método de kmeans con el número de clusters obtenido y graficarlos.
Ruerde utilizar la semilla de random_state=42

```python
kmeans_p4 = KMeans(n_clusters=4, random_state=42).fit(df_escalado_p4)
centroids_p4 = kmeans_p4.cluster_centers_
centroids_p4

labels_p4 = kmeans_p4.predict(df_escalado_p4)
data_p4['label'] = labels_p4
data_p4['cellphone_id'] = id_cell

x = df_escalado_p4['rating'].values
y = df_escalado_p4['age'].values
z = df_escalado_p4['gender'].values

colores=['red','green','blue','orange']
asignar=[]
for row in labels_p4:
     asignar.append(colores[row])
plt.rcParams["figure.figsize"] = (12,7)
ax = plt.axes(projection='3d')
ax.scatter3D(x, y, z, c=asignar)
ax.scatter3D(centroids_p4[:, 0], centroids_p4[:, 1],centroids_p4[:, 2], marker = '*',c=['red','green','blue','orange'],
            s=900)
ax.set_xlabel("Rating")
ax.set_ylabel("Age")
ax.set_zlabel("Gender")
plt.title('Agrupación k-means puesto por el comprador')
plt.show()
```